 Grundlagen der Programmierung in Python

Programmieren besteht aus einer Aneinanderreihung von Befehlen, die der Python Interpreter ausführt. Ein einfacher Befehl ist die Zuweisung eines Wertes zu einer **Variablen**:

```python
x = 1
# Diese Zelle auswählen und mit `<SHIFT> + <ENTER>` ausführen.
# Der letzte Rückgabewert der Zelle wird dann unten ausgegeben,
# hier also der Wert von `x`:
x
```

Output:
```
1
```

> **Hinweis:** Dieses Kursmaterial besteht aus interaktiven [Jupyter Notebooks](http://jupyter.org). Jede _Zelle_, wie die obige, könnt ihr per Klick oder mit den Pfeiltasten auswählen und mit `<SHIFT> + <ENTER>` ausführen. Probiert's mal mit der Zelle oben!
> 
> **Geht, um dem Kurs zu folgen, _alle_ Zellen nacheinander durch und führt sie aus.**
>
> Wie hier demonstriert gibt das Jupyter Notebook immer den letzten Rückgabewert einer Zelle aus. Wir können auch die `print()` Funktion verwenden, um Ausgaben zu generieren:

```python
print(x)
```

Output:
```
1
```

Der **Wert** der Variable `x` ist nun der Integer `1`. Werte haben immer einen **Typ** wie

- `int` für ganze Zahlen, z.B. `1`, `42`, `-10`,
- `float` für Fließkommazahlen, z.B. `0.5`, `3.14`, `1e10`,
- `str` für Zeichenketten (_Strings_), z.B. `"Hello World!"`,
- `boolean` für Wahrheitswerte (_Booleans_), also `True` und `False`

> Die Typen der Python Standardlibrary findet ihr in der [Dokumentation](https://docs.python.org/3/library/stdtypes.html).

Da Variablen in Python nicht typisiert sind, können der gleichen Variable nacheinander Werte verschiedener Typen zugewiesen werden:

```python
type(x)
```

Output:
```
int
```

```python
x = 0.5
type(x)
```

Output:
```
float
```

Weiterhin können Werte in einen anderen Typ konvertiert werden:

```python
x = int(0.5) # Bei der Konvertierung zu int wird immer abgerundet!
x
```

Output:
```
0
```

Die grundlegenden Rechenoperationen `+`, `-`, `*`, `/` und `**` sind ebenfalls in Python verfügbar und verhalten sich, wie man es erwartet:

```python
1 + 3
```

Output:
```
4
```

```python
3 / 2
```

Output:
```
1.5
```

> Die Integer Division ist außerdem in Python 3 als Operator `//` verfügbar:

```python
3 // 2
```

Output:
```
1
```

```python
3**2
```

Output:
```
9
```

## Aufgabe 1 - Werte und Typen

Der Operator zur Potenzierung ist in Python `**`. Weise einer Variablen `y` folgende Werte zu und lasse dir dann ihren Wert und Typ ausgeben. Stimmen Wert und Typ mit deinen Erwartungen überein?

a) $4^3$

> **Hinweis:** Dies ist eine Übungsaufgabe. Verwende die Zelle unten, um sie zu lösen. **Entferne** dazu den Platzhalter-Code
>
> ```python
> # DEINE LÖSUNG HIER
> ```
>
> und schreibe stattdessen den Code zur Lösung der Aufgabe.

<!-- grade_id: 101-1a-sol, points: 0, grade: false, solution: true -->
```python
y = 4**3
y, type(y)
```

Output:
```
(64, int)
```

> **Hinweis:** Auf die Lösung jeder Aufgabe folgt eine Test-Zelle wie die folgende, mit der du deine Lösung überprüfen kannst. Führe einfach diese Zelle aus. Wenn etwas an deiner Lösung nicht stimmt, erhälst du eine Fehlermeldung mit Hinweisen.

<!-- grade_id: 101-1a-test, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
try:
    y
except NameError:
    raise NameError("Es gibt keine Variable 'y'. Weise den Wert einer Variablen mit diesem Namen zu.")
assert_equal(y, 64, "Die Variable hat nicht den richtigen Wert. Überprüfe deine Rechnung.")
print("👍 Sehr gut!")
```

Output:
```
👍 Sehr gut!
```

b) $2+3.4^2$

<!-- grade_id: 101-1b-sol, points: 0, grade: false, solution: true -->
```python
y = 2 + 3.4**2
y, type(y)
```

Output:
```
(13.559999999999999, float)
```

<!-- grade_id: 101-1b-test, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
try:
    y
except NameError:
    raise NameError("Es gibt keine Variable 'y'. Weise den Wert einer Variablen mit diesem Namen zu.")
assert_equal(y, 2 + 3.4**2, "Die Variable hat nicht den richtigen Wert. Überprüfe deine Rechnung.")
print("Ok, du hast es verstanden.")
```

Output:
```
Ok, du hast es verstanden.
```

## Strings

Strings sind **Zeichenketten** wie:

```python
s = "Hello World"
```

und werden in Python vom Typ `str` repräsentiert. Um einen String zu erstellen können wir einzelne (`'`), doppelte (`"`) oder dreifache (`'''` oder `"""`, für mehrzeilige Strings) Anführungszeichen verwenden, sodass das jeweils andere Anführungszeichen im String verwendet werden kann:

```python
s.count(s)
```

Output:
```
1
```

```python
"I'm"
```

Output:
```
"I'm"
```

Alternativ können Steuerzeichen im String auch _escaped_ werden:

```python
"Say \"hi\""
```

Output:
```
'Say "hi"'
```

### Strings sind Reihen

Da Strings eine Aneinanderreihung von Elementen (in diesem Fall Textzeichen) darstellen, können viele Funktionen mit Strings verwendet werden, die mit **Reihen** arbeiten. Dazu gehören:

```python
len(s) # gibt die Zahl der Reihenelemente zurück
```

Output:
```
11
```

```python
s[0] # gibt das Element der Reihe an der Stelle 0 zurück
```

Output:
```
'H'
```

```python
s + "!" # Reihen können kombiniert werden
```

Output:
```
'Hello World!'
```

### Strings sind Objekte

Die meisten "Dinge" in Python sind **Objekte**, d.h. neben ihrem Typ besitzen sie assoziierte **Attribute** und **Methoden**, auf die über einen Punkt `.` zugegriffen werden kann. Neben Strings sind bspw. auch Werte der anderen schon bekannten Datentypen wie `int(5)` und sogar Funktionen und die Datentypen selbst Objekte.

Ein Attribut eines Objekts ist eine Variable, die gelesen und gesetzt werden kann, wie bspw. `x.shape`. Eine Methode ist eine Funktion, die das Objekt zur Verfügung stellt, wie `x.min()`.

Einige Beispiele für Methoden, die Strings zur Verfügung stellen, sind:

```python
s.upper()
```

Output:
```
'HELLO WORLD'
```

```python
s.split()
```

Output:
```
['Hello', 'World']
```

```python
s.index('World')
```

Output:
```
6
```

> **Hinweis:** In Jupyter Notebooks können wir die **`<TAB>`-Vervollständigung** verwenden um die assoziierten Attribute und Methoden eines Objekts zu inspizieren:
>
> ```python
> s = "Hello World"
> # Zelle ausführen, dann:
> s.<TAB> # `s.` tippen und die <TAB>-Taste drücken
> ```
>
> Dies zeigt die verfügbaren Attribute und Methoden des Strings `s` an. Die `<TAB>`-Vervollständigung für eine Variable funktioniert erst nachdem die Variable erstellt wurde, also die Zelle einmal ausgeführt wurde.
>
> Um herauszufinden, was eine Funktion oder Methode tut, könnt ihr im Jupyter Notebook ein Fragezeichen `?` verwenden:
>
>```python
>In [1]: s.split?
>```
>```markdown
>Docstring:
>S.split(sep=None, maxsplit=-1) -> list of strings
>
>Return a list of the words in S, using sep as the
>delimiter string.  If maxsplit is given, at most maxsplit
>splits are done. If sep is not specified or is None, any
>whitespace string is a separator and empty strings are
>removed from the result.
>Type:      builtin_function_or_method
>```
>
> Schreibt ihr stattdessen zwei Fragezeichen `??` zeigt das Jupyter Notebook die gesamte Definition der Funktion oder Methode an.
>
> **Verwendet die `<TAB>`-Vervollständigung und die `?`-Dokumentation häufig um hilfreiche Attribute und Methoden zu finden und zu verstehen!**

## Aufgabe 2 - Strings und Dokumentation

Finde im folgenden String mithilfe einer Methode, wie häufig der Buchstabe `"A"` auftaucht und weise den Wert einer Variable `n` zu. Probiere die `<TAB>`-Vervollständigung und die `?`-Dokumentation aus um eine passende Methode zu finden.

<!-- grade_id: 101-2-sol, points: 0, grade: false, solution: true -->
```python
s = "CAGTACCAAGTGAAAGAT"
### BEGIN SOLUTION
n = s.count("A")
### END SOLUTION
print(n)
```

Output:
```
8
```

<!-- grade_id: 101-2-test, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
try:
    y
except NameError:
    raise NameError("Es gibt keine Variable 'n'. Weise den Wert einer Variablen mit diesem Namen zu.")
assert_equal(n, 8, "Das ist nicht die richtige Anzahl. Versuch's mal mit der `count` Methode!")
print("Klappt!")
```

Output:
```
Klappt!
```

### String-Formatierung

Eine wichtige Methode ist `str.format()`, die markierte Platzhalter im String mit Werten ersetzt, die der Methode übergeben werden:

```python
x = 10
"Der Wert von x ist {}".format(x)
```

Output:
```
'Der Wert von x ist 10'
```

Ein Platzhalter wird durch eine öffnende `{` und schließende `}` geschweifte Klammer begrenzt und kann eine Bezeichnung des Platzhalters sowie Formatierungsoptionen beinhalten:

```python
"{raw_value} ist gerundet {rounded_value:.3f}.".format(raw_value=2/3, rounded_value=2/3)
```

Output:
```
'0.6666666666666666 ist gerundet 0.667.'
```

> **Hinweis:** Die vollständige Syntax der String-Formatierung ist in der [Dokumentation](https://docs.python.org/2/library/string.html#format-string-syntax) einzusehen.

## Aufgabe 3 - String Formatierung

Schreibe deinen Namen in die Variable `name`. Verwende dann die `format` Methode um aus `s` und `name` den Gruß `"Hello World, mein Name ist __DEIN_NAME__!"` zusammenzusetzen. Weise den zusammengesetzten String der Variable `greeting` zu.

<!-- grade_id: 101-3-sol, points: 0, grade: false, solution: true -->
```python
s = "Hello World"
name = "__DEIN_NAME__"
### BEGIN SOLUTION
greeting = s + ", mein Name ist {name}!".format(name=name)
### END SOLUTION
print(greeting)
```

Output:
```
Hello World, mein Name ist __DEIN_NAME__!
```

<!-- grade_id: 101-3-test, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_in
try:
    greeting
except NameError:
    raise NameError("Es gibt keine Variable 'greeting'. Weise den Wert einer Variablen mit diesem Namen zu.")
assert_in("mein Name ist", greeting)
print("Hello {}! 👋".format(name))
```

Output:
```
Hello __DEIN_NAME__! 👋
```

## Listen und Tupel

Während Strings einzelne Zeichen aneinanderreihen, repräsentieren **Listen** und **Tupel** eine Reihe _beliebiger_ Werte. Die Elemente einer Liste können verändert werden, während ein Tupel unveränderlich ist. Wir können Listen erstellen, indem wir die Elemente durch Kommata getrennt in eckigen Klammern `[` und `]` angeben, und Tupel durch die Verwendung runder Klammern `(` und `)`:

```python
l = [ 4, 0.5, "Alice" ]
l
```

Output:
```
[4, 0.5, 'Alice']
```

```python
t = ( "Bob", True )
t
```

Output:
```
('Bob', True)
```

### Auch Listen und Tupel sind Reihen

Wie bei jeder Reihe können wir die Anzahl der Elemente bestimmen und auf einzelne Elemente über ihren Index zugreifen:

```python
len(l), len(t)
```

Output:
```
(3, 2)
```

```python
l[0] # Indexierung beginnt in Python mit dem Index 0
```

Output:
```
4
```

```python
l[1]
```

Output:
```
0.5
```

```python
l[2]
```

Output:
```
'Alice'
```

```python
t[0]
```

Output:
```
'Bob'
```

```python
t[1]
```

Output:
```
True
```

### Listen sind veränderlich

Anders als Strings und Tupel können Listen jedoch verändert werden, indem Elemente verändert, hinzugefügt oder entfernt werden:

```python
l[1] = -2.2 # Weise der Liste einen neuen Wert beim Index 1 zu
l
```

Output:
```
[4, -2.2, 'Alice']
```

```python
l.append(-3) # Füge einen neuen Wert am Ende der liste hinzu
l
```

Output:
```
[4, -2.2, 'Alice', -3]
```

```python
l.pop() # Entfernt das letzte Element (bzw. das am angegebenen Index) aus der Liste und gibt es zurück
```

Output:
```
-3
```

## Slicing

Du kannst mit der _Slicing_ Syntax auf Teile einer Reihe zugreifen:

```python
slice = list[start:stop:step]
```

Dabei bezeichnet `start` den ersten Index und `stop` den letzten nicht mehr enthaltenen Index des Abschnitts. Mit `step` kannst du eine Schrittweite angeben, in der die Reihe iteriert werden soll.

Du musst nicht alle drei Argumente angeben. Per default ist dann `start=0`, `stop=len(list)` und `step=1`:

```python
l[:2] # Der erste Index ist per default start=0
```

Output:
```
[4, -2.2]
```

```python
l[2:] # Der letzte Index ist per default das Ende der Reihe
```

Output:
```
['Alice']
```

```python
l[::2] # Jedes zweite Element
```

Output:
```
[4, 'Alice']
```

```python
l[::-1] # Umgekehrte Reihenfolge
```

Output:
```
['Alice', -2.2, 4]
```

> **Hinweis:** Slicing ist ein mächtiges Werkzeug um kurzen, prägnanten Code zu schreiben und wird dir bei der Arbeit mit _Numpy Arrays_ noch sehr häufig begegnen.

## Dictionaries

Ein weiterer wichtiger Datentyp ist das **Dictionary**. Ein Dictionary ordnet jeweils einem _Key_ einen _Value_ zu und wird durch geschweifte Klammern `{` und `}` erstellt:

```python
d = { 'a':1, 'b':2, 'c':3 }
d
```

Output:
```
{'a': 1, 'b': 2, 'c': 3}
```

Auf einzelne Werte kann über ihren Key zugegriffen werden:

```python
d['a']
```

Output:
```
1
```

Dictionaries sind veränderlich wie Listen:

```python
d['d'] = 4
d
```

Output:
```
{'a': 1, 'b': 2, 'c': 3, 'd': 4}
```

Es müssen nicht unbedingt Strings als Keys verwendet werden, und auch verschiedene Datentypen sind möglich:

```python
e = { 'some_key': 4.2, 3517: 'some_value' }
e
```

Output:
```
{'some_key': 4.2, 3517: 'some_value'}
```

```python
e['some_key']
```

Output:
```
4.2
```

```python
e[3517]
```

Output:
```
'some_value'
```

## Aufgabe 4 - Dictionaries

Schreibe ein Wörterbuch mit einigen deutschen Wörtern und ihrer Übersetzung ins Englische. Weise es der Variable `d` zu und verwende es dann, indem du dir einige Werte darin ausgeben lässt.

<!-- grade_id: 101-4-sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
d = {
    "Wörterbuch": "dictionary"
}
### END SOLUTION
d["Wörterbuch"]
```

Output:
```
'dictionary'
```

<!-- grade_id: 101-4-test, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_true
try:
    d
except NameError:
    raise NameError("Es gibt keine Variable 'd'. Weise den Wert einer Variablen mit diesem Namen zu.")
assert_true(len(d) > 0, "Das Wörterbuch hat keine Einträge.")
print("Wow, {} Einträge im Wörterbuch!".format(len(d)))
```

Output:
```
Wow, 1 Einträge im Wörterbuch!
```